The role of GABA in semantic memory and its neuroplasticity

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

This study examined the changes in ATL GABA levels induced by cTBS and its relationship with BOLD signal changes and performance in a semantic task. The findings suggest that the increase in ATL GABA levels induced by cTBS is associated with a decrease in BOLD signal. The relationship between ATL GABA levels and semantic task performance is nonlinear, and more specifically, the authors propose that the relationship is an inverted U-shaped relationship.

Strengths:

The findings of the research regarding the increase of GABA and decrease of BOLD caused by cTBS, as well as the correlation between the two, appear to be reliable. This should be valuable for understanding the biological effects of cTBS.

Weakness:

I am pleased to see the authors' feedback on my previous questions and suggestions, and I believe the additional data analysis they have added is helpful. Here are my reserved concerns and newly discovered issues.

(1) Regarding the Inverted U-Shaped Curve In the revised manuscript, the authors have accepted some of my suggestions and conducted further analysis, which is now presented in Figure 3B. These results provide partial support for the authors' hypothesis. However, I still believe that the data from this study hardly convincingly support an inverted U-shaped distribution relationship.

The authors stated in their response, "it is challenging to determine the optimal level of ATL GABA," but I think this is achievable. From Figures 4C and 4D, the ATL GABA levels corresponding to the peak of the inverted U-shaped curve fall between 85 and 90. In my understanding, this can be considered as the optimal level of ATL GABA estimated based on the existing data and the inverted U-shaped curve relationship. However, in the latter half of the inverted U-shaped curve, there are quite few data points, and such a small number of data points hardly provides reliable support for the quantitative relationship in the latter half of the curve. I suggest that the authors should at least explicitly acknowledge this and be cautious in drawing conclusions. I also suggest that the authors consider fitting the data with more types of non-linear relationships, such as a ceiling effect (a combination of a slope and a horizontal line), or a logarithmic curve.

We appreciate R1’s comments. Inverted U-shaped relationships are well-established in neuroscience, particularly in the context of neurotransmitter concentrations (e.g., dopamine, acetylcholine, noradrenaline) and their influence on cognitive functions such as working memory and cognitive control (Aston-Jones & Cohen., 2005; Cools & D'Esposito., 2011; Vijayraghavan et al., 2007; He & Zempel., 2013). Recently, Ferri et al. (2017) demonstrated an inverted U-shaped relationship between excitation-inhibition balance (EIB: the ratio of Glx and GABA) and multisensory integration, showing that both excessive and insufficient inhibition negatively impact functionality. Given that GABA is the brain’s primary inhibitory neurotransmitter, our findings suggest that ATL GABA may play a similar regulatory role in semantic memory function.

While our statistical modelling approach demonstrated that the inverted U-shaped function was the best-fitting model for our current data in explaining the relationship between ATL GABA and semantic memory, we acknowledge the limitation of having fewer data points in the latter half (right side) of the curve, where excessive ATL GABA levels are associated with poorer semantic performance. Following R1’s suggestion, we have explicitly acknowledged this limitation in the revised manuscript and exercised caution in our discussion.

Discussion, p.17, line 408

"However, our findings should be interpreted with caution due to the limitation of having fewer data points in the latter half (right side) of the inverted U-shaped curve. Future studies incorporating GABA agonists could help further validate and refine these findings."

Following R1’s latter suggestion, we tested a logarithmic curve model. The results showed significant relationships between ATL GABA and semantic performance (R2 = 0.544, p < 0.001) and between cTBS-induced changes in ATL GABA and semantic performance (R2 = 0.202, p < 0.001). However, the quadratic (inverted U-shaped) model explained more variance than the logarithmic model, as indicated by a higher R2 and lower BIC. Model comparisons further confirmed that the inverted U-shaped model provided the best fit for both ATL GABA in relation to semantic performance (Fig. 4C) and cTBS-induced ATL GABA changes in relation to semantic function (Fig. 4D).

(2) In Figure 2F, the authors demonstrated a strong practice effect in this study, which to some extent offsets the decrease in behavioral performance caused by cTBS. Therefore, I recommend that the authors give sufficient consideration to the practice effect in the data analysis.

One issue is the impact of the practice effect on the classification of responders and non-responders. Currently, most participants are classified as non-responders, suggesting that the majority of the population may not respond to the cTBS used in this study. This greatly challenges the generalizability of the experimental conclusions. However, the emergence of so many non-responders is likely due to the prominent practice effect, which offsets part of the experimental effect. If the practice effect is excluded, the number of responders may increase. The authors might estimate the practice effect based on the vertex simulation condition and reclassify participants after excluding the influence of the practice effect.

Another issue is that considering the significant practice effect, the analysis in Figure 4D, which mixes pre- and post-test data, may not be reliable.

We appreciate Reviewer 1’s thoughtful comments regarding the practice effect and its potential impact on our findings. Our previous analysis revealed a strong practice effect on reaction time (RT), with participants performing tasks faster in the POST session, regardless of task condition (Fig. S3). Given our hypothesis that inhibitory ATL cTBS would disrupt semantic task performance, we accounted for this by using inverse efficiency (IE), which combines accuracy and RT. This analysis demonstrated that ATL cTBS disrupted semantic task performance compared to both control stimulation (vertex) and control tasks, despite the practice effect (i.e., faster RT in the POST session), thereby supporting our hypothesis. These findings may suggest that the effects of ATL cTBS were more subtly reflected in semantic task accuracy rather than RT.

Regarding inter-individual variability in response to rTMS/TBS, prior studies have shown that 50–70% of participants are non-responders, either do not respond or respond in an unexpected manner (Goldsworthy et al., 2014; Hamada et al., 2013; Hinder et al., 2014; Lopez-Alonso et al., 2014; Maeda et al., 2000a; Müller-Dahlhaus et al., 2008). Our previous study (Jung et al., 2022) using the same semantic task and cTBS protocol was the first to explore TBS-responsiveness variability in semantic memory, where 12 out of 20 participants (60%) were classified as responders. The proportion of responders and non-responders in the current study aligns with previous findings, suggesting that this variability is expected in TBS research.

However, we acknowledge R1’s concern that the strong practice effect may have influenced responder classification. To address this, we estimated the practice effect using the vertex stimulation condition and reclassified participants accordingly by adjusting ATL stimulation performance (IE) relative to vertex stimulation performance (IE). This reclassification identified nine responders (an increase of two), aligning with the typical responder proportion (52%) reported in the TBS literature. Overall, we replicated the previous findings with improved statistical robustness.

A 2×2×2 ANOVA was conducted with task (semantic vs. control) and session (PRE vs. POST) as within-subject factors, and group (responders vs. non-responders) as a between-subject factor. The analysis revealed a significant interaction between the session and group (F1, 15 = 10.367, p = 0.006), a marginally significant interaction between the session and task (F1, 15 = 4.370, p = 0.054), and a significant 3-way interaction between the session, task, and group (F1, 15 = 7.580, p = 0.015). Post hoc t-tests showed a significant group difference in semantic task performance following ATL stimulation (t = 2.349, p = 0.033). Post hoc paired t-test demonstrated that responders exhibited poorer semantic task performance following the ATL cTBS (t = -5.281, p < 0.001), whereas non-responders showed a significant improvement (t = 3.206, p = 0.007) (see Figure. 3A).

Notably, no differences were observed between responders and non-responders in the control task performance across pre- and post-stimulation sessions, confirming that the practice effect was successfully controlled (Figure. 3B).

We performed a 2 x 2 ANOVA with session (pre vs. post) as a within subject factor and with group (responders vs. non-responders) as a between subject factor to examine the effects of group in ATL GABA levels. The results revealed a significant main effect of session (F1, 14 = 39.906, p < 0.001) and group (F1, 14 = 9.677, p = 0.008). Post hoc paired t-tests on ATL GABA levels showed a significant increase in regional ATL GABA levels following ATL stimulation for both responders (t = -3.885, p = 0.002) and non-responders (t = -4.831, p = 0.001). Furthermore, we replicated our previous finding that baseline GABA levels were significantly higher in responders compared to non-responders (t = 2.816, p = 0.007) (Figure. 3C). This pattern persisted in the post-stimulation session (t = 2.555, p = 0.011) (Figure. 3C).

Accordingly, we have revised the Methods and Materials (p 26, line 619), Results (p11, line 233-261), and Figure 3.

(3) The analysis in Figure 3A has a double dipping issue. Suppose we generate 100 pairs of random numbers as pre- and post-test scores, and then group the data based on whether the scores decrease or increase; the pre-test scores of the group with decreased scores will have a very high probability of being higher than those of the group with increased scores. Therefore, the findings in Figure 3A seem to be meaningless.

Yes, we agreed with R1’s comments. However, Figure 3A illustrates interindividual responsiveness patterns, while Figure 3B demonstrates that these results account for practice effects, incorporating new analyses.

(4) The authors use IE as a behavioral measure in some analyses and use accuracy in others. I recommend that the authors adopt a consistent behavioral measure.

We appreciate Reviewer 1’s suggestion. In examining the relationship between ATL GABA and semantic task performance, we have found that only semantic accuracy—not reaction time (RT) or inverse efficiency (IE)—shows a significant positive correlation and regression with ATL GABA levels and semantic task-induced ATL activation, both in our previous study (Jung et al., 2017) and in the current study. ATL GABA levels were not correlated with semantic RT (Jung et al., 2017: r = 0.34, p = 0.14, current study: r = 0.26, p = 0.14). It should be noted that there were no significant correlations between ATL GABA levels and semantic inverse efficiency (IE) in both studies (Jung et al., 2017: r = 0.13, p = 0.62, current study: r = 0.22, p = 0.44). As a result, we found no significant linear and non-linear relationship between ATL GABA levels and RT (linear function R2 = 0.21, p =0.45, quadratic function: R2 = 0.17, p = 0.21) and between ATL GABA levels and IE (linear function R2 = 0.24, p =0.07, quadratic function: R2 = 2.24, p = 0.12).

The absence of a meaningful relationship between ATL GABA and semantic RT or IE may be due to the following reasons: 1) RT is primarily associated with premotor and motor activation during semantic processing rather than ATL activation; 2) ATL GABA is likely to play a key role in refining distributed semantic representations through lateral inhibition, which sharpens the activated representation (Jung et al., 2017; Liu et al. 2011; Isaacson & Scanziani., 2011). This sharpening process may contribute to more accurate semantic performance (Jung et al., 2017). In our semantic task, for example, when encountering a camel (Fig. 1B), multiple semantic features (e.g., animal, brown, desert, sand, etc.) are activated. To correctly identify the most relevant concept (cactus), irrelevant associations (tree) must be suppressed—a process that likely relies on inhibitory mechanisms. Given this theoretical framework, we have used accuracy as the primary measure of semantic performance to elucidate the ATL GABA function.

Reviewer #2 (Public review):

Summary:

The authors combined inhibitory neurostimulation (continuous theta-burst stimulation, cTBS) with subsequent MRI measurements to investigate the impact of inhibition of the left anterior temporal lobe (ATL) on task-related activity and performance during a semantic task and link stimulation-induced changes to the neurochemical level by including MR spectroscopy (MRS). cTBS effects in the ATL were compared with a control site in the vertex. The authors found that relative to stimulation of the vertex, cTBS significantly increased the local GABA concentration in the ATL. cTBS also decreased task-related semantic activity in the ATL and potentially delayed semantic task performance by hindering a practice effect from pre to post. Finally, pooled data with their previous MRS study suggest an inverted u-shape between GABA concentration and behavioral performance. These results help to better understand the neuromodulatory effects of non-invasive brain stimulation on task performance.

Strengths:

Multimodal assessment of neurostimulation effects on the behavioral, neurochemical, and neural levels. In particular, the link between GABA modulation and behavior is timely and potentially interesting.

Weaknesses:

The analyses are not sound. Some of the effects are very weak and not all conclusions are supported by the data since some of the comparisons are not justified. There is some redundancy with a previous paper by the same authors, so the novelty and contribution to the field are overall limited. A network approach might help here.

Reviewer #3 (Public review):

Summary:

The authors used cTBS TMS, magnetic resonance spectroscopy (MRS), and functional magnetic resonance imaging (fMRI) as the main methods of investigation. Their data show that cTBS modulates GABA concentration and task-dependent BOLD in the ATL, whereby greater GABA increase following ATL cTBS showed greater reductions in BOLD changes in ATL. This effect was also reflected in the performance of the behavioural task response times, which did not subsume to practice effects after AL cTBS as opposed to the associated control site and control task. This is in line with their first hypothesis. The data further indicates that regional GABA concentrations in the ATL play a crucial role in semantic memory because individuals with higher (but not excessive) GABA concentrations in the ATLs performed better on the semantic task. This is in line with their second prediction. Finally, the authors conducted additional analyses to explore the mechanistic link between ATL inhibitory GABAergic action and semantic task performance. They show that this link is best captured by an inverted U-shaped function as a result of a quadratic linear regression model. Fitting this model to their data indicates that increasing GABA levels led to better task performance as long as they were not excessively low or excessively high. This was first tested as a relationship between GABA levels in the ATL and semantic task performance; then the same analyses were performed on the pre and post-cTBS TMS stimulation data, showing the same pattern. These results are in line with the conclusions of the authors.

Comments on revisions:

The authors have comprehensively addressed my comments from the first round of review, and I consider most of their answers and the steps they have taken satisfactorily. Their insights prompted me to reflect further on my own knowledge and thinking regarding the ATL function.

I do, however, have an additional and hopefully constructive comment regarding the point made about the study focusing on the left instead of bilateral ATL. I appreciate the methodological complexities and the pragmatic reasons underlying this decision. Nevertheless, briefly incorporating the justification for this decision into the manuscript would have been beneficial for clarity and completeness. The presented argument follows an interesting logic; however, despite strong previous evidence supporting it, the approach remains based on an assumption. Given that the authors now provide the group-level fMRI results captured more comprehensively in Supplementary Figure 2, where the bilateral pattern of fMRI activation can be observed in the current data, the authors could have strengthened their argument by asserting that the activation related to the given semantic association task in this data was bilateral. This would imply that the TMS effects and associated changes in GABA should be similar for both sites. Furthermore, it is worth noting the approach taken by Pobric et al. (2007, PNAS), who stimulated a site located 10 mm posterior to the tip of the left temporal pole along the middle temporal gyrus (MTG) and not the bilateral ATL.

We appreciate the reviewer’s constructive comment regarding the focus on the left ATL rather than bilateral ATL in our study. Accordingly, we have added the following paragraph in the Supplementary Information.

“Justification of target site selection and cTBS effects

Evidence suggests that bilateral ATL systems contribute to semantic representation (for a review, see Lambon Ralph., 2017). Consistent with this, our semantic task induced bilateral ATL activation (Fig. S2). Thus, stimulating both left and right ATL could provide a more comprehensive understanding of cTBS effects and its GABAergic function.

Previous rTMS studies have applied inhibitory stimulation to the left vs. right ATL, demonstrating that stimulation at either site significantly disrupted semantic task performance (Pobric et al., 2007, PNAS; Pobric et al., 2010, Neuropsychologia; Lambon Ralph et al., 2009, Cerebral Cortex). Importantly, these studies reported no significant difference in rTMS effects between left and right ATL stimulation, suggesting that stimulating either hemisphere produces comparable effects on semantic processing. In the current study, we combined cTBS with multimodal imaging to investigate its effects on the ATL. Given our study design constraints (including the need for a control site, control task, and control stimulation) and limitations in scanning time, we selected the left ATL as the target region. This choice also aligned with the MRS voxel placement used in our previous study (Jung et al., 2017), allowing us to combine datasets and further investigate GABAergic function in the ATL. Accordingly, cTBS was applied to the peak coordinate of the left ventromedial ATL (MNI -36, -15, -30) as identified by previous fMRI studies (Binney et al., 2010; Visser et al., 2012).

Given that TMS pulses typically penetrate 2–4 cm, we acknowledge the challenge of reaching deeper ventromedial ATL regions. However, our findings indicate that cTBS effectively modulated ATL function, as evidenced by reduced task-induced regional activity, increased ATL GABA concentrations, and poorer semantic performance, confirming that TMS pulses successfully influenced the target region. To further validate these effects, we conducted an ROI analysis centred on the ventromedial ATL (MNI -36, -15, -30), which revealed a significant reduction in ATL activity during semantic processing following ATL stimulation (t = -2.43, p = 0.014) (Fig. S7). This confirms that cTBS successfully modulated ATL activity at the intended target coordinate.”

We appreciate R3's comment regarding the approach taken by Pobric et al. (2007, PNAS), who stimulated a site 10 mm posterior to the tip of the left temporal pole along the middle temporal gyrus (MTG). This approach has been explicitly discussed in our previous papers and reviews (e.g., Lambon Ralph, 2014, Proc. Royal Society B). Our earlier use of lateral ATL stimulation at this location (Pobric et al. 2007; Lambon Ralph et al. 2009; Pobric et al. 2010) was based on its alignment with the broader ATL region commonly atrophied in semantic dementia (cf. Binney et al., 2010 for a direct comparison of SD atrophy, fMRI data and the TMS region). Since these original ATL TMS investigations, a series of distortion-corrected or distortion-avoiding fMRI studies (e.g., Binney et al 2010; Visser et al, various, Hoffman et al., various; Jackson et al., 2015) have demonstrated graded activation differences across the ATL. While weaker activation is present at the original lateral ATL (MTG) stimulation site, the peak activation is maximal in the ventromedial ATL—a finding that was also observed in the current study. Accordingly, we selected the ventromedial ATL as our target site for stimulation.

Following these points, we have revised the manuscript in the Methods and Materials.

Transcranial magnetic stimulation p23, line 525-532,

“Previous rTMS studies targeted a lateral ATL site 10 mm posterior to the temporal pole on the middle temporal gyrus (MTG) (Pobric et al. 2007; Lambon Ralph et al. 2009; Pobric et al. 2010), aligning with the broader ATL region typically atrophied in semantic dementia (Binney et al. 2010). However, distortion-corrected fMRI studies (Binney et al. 2010; Visser et al. 2012) have revealed graded activation differences across the ATL, with peak activation in the ventromedial ATL. Based on these findings, we selected the target site in the left ATL (MNI -36, -15, -30) from a prior distortion-corrected fMRI study (Binney et al. 2010; Visser et al. 2012 that employed the same tasks as our study (for further details, see the Supplementary Information).”

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

The authors have responded to all my comments and I found most of the responses reasonable and sufficient. However, I have one remaining point: I pointed out before that the scope of this paper is somehow narrow and asked for a network analysis. I found the response to my question somehow puzzling since the authors write:

"However, it is important to note that we did not find any significant correlations between ATL GABA changes and cTBS-induced changes in the functional connectivity. Consequently, we are currently preparing another paper that specifically addresses the network-level changes induced by ATL cTBS."

I don't understand the logic here. Even in the absence of significant correlations between ATL GABA changes and cTBS-induced changes in connectivity, it would be interesting to know how baseline connectivity is correlated with the induced changes. I am not sure if it is adequate to squeeze another paper out of the dataset instead of reporting it here as suggested.

We apologise that our previous response was not clear. To examine cTBS-induced network-level changes, we conducted ROI analyses targeting key semantic regions, including the bilateral ATL, inferior frontal gyrus (IFG), and posterior middle temporal gyrus (pMTG), as well as Psychophysiological Interactions (PPI) using the left ATL as a seed region. The ROI analysis revealed that ATL stimulation significantly decreased task-induced activity in the left ATL (target region) while increasing activity in the right ATL and left IFG. PPI analyses showed that ATL stimulation enhanced connectivity between the left ATL and the right ATL (both ventromedial and lateral ATL), bilateral IFG, and bilateral pMTG, suggesting that ATL stimulation modulates a bilateral semantic network.

Building on these findings, we conducted Dynamic Causal Modeling (DCM) to estimate and infer interactions among predefined brain regions across different experimental conditions (Friston et al., 2003). The bilateral ventromedial ATL, lateral ATL, IFG, and pMTG were defined as network nodes with mutual connections. Our model examined cTBS effects at the left ATL under both baseline (intrinsic) and semantic task (modulatory) conditions, estimating 56 intrinsic parameters for baseline connectivity and testing 16 different modulatory models to assess cTBS-induced connectivity changes during semantic processing. Here, we briefly summarize the key DCM analysis results: 1) ATL cTBS significantly altered effective connectivity between the left and right lateral and ventromedial ATL in both intrinsic and modulatory conditions; 2) cTBS increased modulatory connectivity from the right to the left ATL compared to vertex stimulation.

Given the complexity and depth of these findings, we believe that a dedicated paper focusing on the network-level effects of ATL cTBS is necessary to provide a more comprehensive and detailed analysis, which extends beyond the scope of the current study. It should be noted that no significant relationship was found between ATL GABA levels and ATL connectivity in both PPI and DCM analyses.

Reviewer #3 (Recommendations for the authors):

In response to my comment about the ATL activation being rather medial in the fMRI data and my concern about the TMS pulse perhaps not reaching this site, the authors offer an excellent solution to demonstrate TMS effects to such a medial ATL coordinate. I think that the analyses and figures they provide as a response to this comment and a brief explanation of this result should be incorporated into supplementary materials for methodologically oriented readers. Also, perhaps it would be beneficial to discuss that the effect of TMS on vATL remains a matter of further research to see not just if but also how TMS pulse reaches target coordinates, given the problematic anatomical location of the region.

We appreciate R3’s suggestion. Please, see our reply above.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: A fundamental aspect of neuroscience is understanding neural functioning and plasticity of the brain. The anterior temporal lobe (ATL) is a hub for semantic memory, which generates coherent semantic representations about the world. GABAergic inhibition plays a crucial role in shaping human cognition and plasticity, but it is unclear how this inhibition relates to human semantic memory and its plasticity. Here, we employed a combination of continuous theta burst stimulation (cTBS), MR spectroscopy and fMRI to investigate the role of GABA in semantic memory and its neuroplasticity. Our results demonstrated that the inhibitory cTBS increased regional GABA levels in the ATL and decreased ATL blood-oxygen level-dependent (BOLD) activity during semantic processing. Importantly, changes in GABA levels were strongly associated with changes in regional activity induced by cTBS. These results suggest that GABAergic activity may be the mechanism by which cTBS induces after effects on cortical excitability. Furthermore, individuals with better semantic performance exhibited selective activity in the ATL, attributable to higher concentrations of inhibitory GABA, which can sharpen distributed semantic representations, leading to more precise semantic processing. Our results revealed a non-linear, inverted-U-shape relationship between GABA levels in the ATL and semantic performance, thus offering an explanation for the individual differences in the cTBS effect on task performance. These results provide neurochemical and anatomical specificity in shaping task-related cortical activity and behaviour. Understanding the link between neurochemistry and semantic memory has important implications for understanding individual differences in semantic behaviour and developing therapeutic interventions for patients with semantic impairments.

Revised version: A fundamental aspect of neuroscience is understanding neural functioning and plasticity of the brain. The anterior temporal lobe (ATL) is a hub for semantic memory, which generates coherent semantic representations about the world. GABAergic inhibition plays a crucial role in shaping human cognition and plasticity, but it is unclear how this inhibition relates to human semantic memory and its plasticity. Here, we employed a combination of continuous theta burst stimulation (cTBS), MR spectroscopy and fMRI to investigate the role of GABA in semantic memory and its neuroplasticity. Our results demonstrated that the inhibitory cTBS increased regional GABA levels in the ATL and decreased ATL blood-oxygen level-dependent (BOLD) activity during semantic processing. Importantly, changes in GABA levels were strongly associated with changes in regional activity induced by cTBS. These results suggest that GABAergic activity may be the mechanism by which cTBS induces aftereffects on cortical excitability. Furthermore, individuals with better semantic performance exhibited selective activity in the ATL, attributable to higher concentrations of inhibitory GABA, which can sharpen distributed semantic representations, leading to more precise semantic processing. Our results revealed a non-linear, inverted-U-shape relationship between GABA levels in the ATL and semantic performance, thus offering an explanation for the individual differences in the cTBS effect on task performance. These results provide neurochemical and anatomical specificity in shaping task-related cortical activity and behaviour. Understanding the link between neurochemistry and semantic memory has important implications for understanding individual differences in semantic behaviour and developing therapeutic interventions for patients with semantic impairments.

Added in the revised version: Main

Removed in the revised version: Introduction

Understanding how the brain functions to drive flexible human behaviour has been a fundamental challenge in cognitive neuroscience. The ability to (re)shape our behaviours based on our experiences relies on a flexible mechanism in the brain, which is achieved through the regulation of neural excitation and inhibition (1). In particular, an imbalance between excitatory and inhibitory processes has been associated with various cognitive impairments in several psychiatric disorders (2) such as autism spectrum disorder (3) and schizophrenia (4). Of particular interest is the role of the neurotransmitters, gamma-aminobutyric acid (GABA) and glutamate in coordinating neural functions supporting performance in various cognitive domains. GABA is the primary inhibitory neurotransmitter in the brain, while glutamate is the primary excitatory neurotransmitter. The balance between GABAergic and glutamatergic neurotransmission is crucial for the proper functioning of the brain and the maintenance of optimal behavioural responses in both healthy and diseased states (5-7). While GABA and glutamate are associated with various cognitive processing, the mechanistic link from the neurotransmitters to human cognition is not well understood.

Research has shown that GABAergic inhibition plays a crucial role in various processes, such as sensory processing, attention, memory and learning (5, 6, 8). Furthermore, GABAergic neurotransmission can regulate synaptic plasticity, leading to long-term potentiation (LTP) and long-term depression (LTD) by modulating the activity of excitatory neurons (9-11). Understanding GABAergic inhibition is crucial in uncovering the neurochemical mechanisms underlying human cognition and its neuroplasticity. Research has revealed a link between variability in the levels of GABA in the human brain and individual differences in cognitive behaviour (for a reveiw, see 5). Specifically, GABA levels in the sensorimotor cortex were found to predict individual performance in the related tasks: higher GABA levels were correlated with a slower reaction time in simple motor tasks (12) as well as improved motor control (13) and sensory discrimination (14, 15). Visual cortex GABA concentrations were positively correlated with a stronger orientation illusion (16), a prolonged binocular rivalry (17), while displaying a negative correlation with motion suppression (17). Individuals with greater frontal GABA concentrations demonstrated enhanced working memory capacity (18, 19). Studies on learning have reported the importance of GABAergic changes in the motor cortex for motor and perceptual learning: individuals showing bigger decreases in local GABA concentration can facilitate this plasticity more effectively (12, 20-22). However, the relationship between GABAergic inhibition and higher cognition in humans remains unclear. The aim of the study was to investigate the role of GABA in relation to human higher cognition – semantic memory and its neuroplasticity at individual level.

Earlier version: Semantic memory is a crucial aspect of human cognition, encompassing our knowledge of concepts and meaning, such as words, people, objects and emotion (24, 25). Accumulating and converging evidence indicates that the anterior temporal lobe (ATL) is a transmodal and transtemporal hub of semantic memory that generates coherent semantic representations though interactions with modality-specific brain regions and integration over time/episodes (25-27). The initial and strong evidence supporting this hypothesis comes from semantic dementia patients who show selective semantic degradation in both verbal and non-verbal domains due to progressive ATL-centred atrophy (24, 28-30). Recent studies also have supported this hypothesis using intracranial recordings and cortical stimulation (31, 32), magnetoencephalography (33, 34) and functional magnetic resonance imaging (fMRI) (35-38). Transcranial magnetic stimulation (TMS) studies have further established the link between ATL and semantic memory. Perturbing the ATL with inhibitory repetitive TMS (rTMS) and theta burst stimulation (TBS) resulted in healthy individuals, showing slower reaction time during semantic processing (39-45). Our investigations combining rTMS/TBS with fMRI have revealed the critical role of the ATL in the neuroplasticity of the semantic system, demonstrating the flexible and adaptive nature of the neural mechanisms underpinning semantic memory function (39, 46, 47). Despite the compelling and consistent findings regarding the involvement of the ATL in semantic memory and its capacity for neuroplasticity, the specific ways in which the underlying neurotransmitter systems influence ATL function in semantic memory and its neuroplasticity remain unclear.

Revised version: Semantic memory is a crucial aspect of human cognition, encompassing our knowledge of concepts and meaning, such as words, people, objects and emotion (23, 24). Accumulating and converging evidence indicates that the anterior temporal lobe (ATL) is a transmodal and transtemporal hub of semantic memory that generates coherent semantic representations though interactions with modality-specific brain regions and integration over time/episodes (24-26). The initial and strong evidence supporting this hypothesis comes from semantic dementia patients who show selective semantic degradation in both verbal and non-verbal domains due to progressive ATL-centred atrophy (23, 27-29). Recent studies also have supported this hypothesis using intracranial recordings and cortical stimulation (30, 31), magnetoencephalography (32, 33) and functional magnetic resonance imaging (fMRI) (34-37). Transcranial magnetic stimulation (TMS) studies have further established the link between ATL and semantic memory. Perturbing the ATL with inhibitory repetitive TMS (rTMS) and theta burst stimulation (TBS) resulted in healthy individuals showing slower reaction time during semantic processing (38-44). Our investigations combining rTMS/TBS with fMRI have revealed the critical role of the ATL in the neuroplasticity of the semantic system, demonstrating the flexible and adaptive nature of the neural mechanisms underpinning semantic memory function (38, 45, 46). Despite the compelling and consistent findings regarding the involvement of the ATL in semantic memory and its capacity for neuroplasticity, the specific ways in which the underlying neurotransmitter systems influence ATL function in semantic memory and its neuroplasticity remain unclear.

Earlier version: Previously, we explored neurotransmitter systems on the functioning of the anterior temporal lobe (ATL) in semantic memory using a combination of magnetic resonance spectroscopy (MRS) and fMRI (48). By utilizing MRS, a non-invasive method for measuring neurometabolites such as GABA and glutamate in vivo, we were able to detect and quantify regional GABA and glutamate concentrations in the ATL (49). The concentration of GABA in the ATL showed a positive correlation with performance in semantic tasks and was negatively associated with blood-oxygen level-dependent (BOLD) signal changes during semantic processing. Our results highlighted the critical involvement of GABAergic inhibition in the modulation of neural activity and behaviour related to semantic processing in the ATL. Subsequently, we explored the relationship between regional GABA levels in the ATL and cTBS-induced plasticity in semantic memory (8). To achieve this, we acquired the ATL MRS and fMRI prior to stimulation and then delivered cTBS, an inhibitory protocol (50), to the ATL. We examined how baseline GABA levels in the ATL were associated with changes in semantic task performance after cTBS. The results showed that individuals with higher GABA levels in the ATL exhibited stronger cTBS effects on semantic processing, especially those who displayed inhibitory responses after cTBS. These findings suggest that the GABAergic action in the ATL plays a crucial role in cTBS-induced plasticity in semantic memory, predicting inter-individual variability of cTBS responsiveness.

Revised version: Previously, we explored neurotransmitter systems on the functioning of the anterior temporal lobe (ATL) in semantic memory using a combination of magnetic resonance spectroscopy (MRS) and fMRI (47). By utilizing MRS, a non-invasive method for measuring neurometabolites such as GABA and glutamate in vivo, we were able to detect and quantify regional GABA and glutamate concentrations in the ATL (48). The concentration of GABA in the ATL showed a positive correlation with performance in semantic tasks and was negatively associated with blood-oxygen level-dependent (BOLD) signal changes during semantic processing. Our results highlighted the critical involvement of GABAergic inhibition in the modulation of neural activity and behaviour related to semantic processing in the ATL. Subsequently, we explored the relationship between regional GABA levels in the ATL and cTBS-induced plasticity in semantic memory (8). To achieve this, we acquired the ATL MRS and fMRI prior to stimulation and then delivered cTBS, an inhibitory protocol (49), to the ATL. We examined how baseline GABA levels in the ATL were associated with changes in semantic task performance after cTBS. The results showed that individuals with higher GABA levels in the ATL exhibited stronger cTBS effects on semantic processing, especially those who displayed inhibitory responses after cTBS. These findings suggest that the GABAergic action in the ATL plays a crucial role in cTBS-induced plasticity in semantic memory, predicting inter-individual variability of cTBS responsiveness.

Earlier version: Based on these findings, we hypothesized that GABAergic inhibition in the ATL can affect neural dynamics of the ATL underpinning semantic memory and its neuroplasticity. The study aimed to investigate the neural mechanisms underlying cTBS-induced neuroplasticity in semantic memory by linking cortical neurochemical profiles, task-induced regional activity, and variability in semantic memory capability within the ATL. We used a combination of cTBS, MRS and fMRI to examine the relationship between changes in GABA levels, cortical activity during a semantic task, and semantic task performance. First, we hypothesized that the inhibitory cTBS would increase GABA concentrations and decrease task-induced BOLD signal changes in the ATL. Second, the effects of cTBS on semantic processing can be attributed to GABAergic action in the ATL at the individual level: greater changes in GABA concentrations would result in increasing changes in task-induced BOLD signal during semantic processing and semantic task performance. Furthermore, to address and explore the relationship between regional GABA levels in the ATL and semantic memory function, we combined data from our previous study (48) with the current study’s data. We then explored the function of GABA in the ATL in relation to semantic function. Finally, we extended this GABAergic function to semantic neuroplasticity driven by cTBS.

Revised version: Based on these findings, we hypothesized that GABAergic inhibition in the ATL can affect neural dynamics of the ATL underpinning semantic memory and its neuroplasticity. The study investigated the neural mechanisms underlying cTBS-induced neuroplasticity in semantic memory by linking cortical neurochemical profiles, task-induced regional activity, and variability in semantic memory capability within the ATL. We used a combination of cTBS, MRS and fMRI to examine the relationship between changes in GABA levels, cortical activity during a semantic task, and semantic task performance. First, we hypothesized that the inhibitory cTBS would increase GABA concentrations and decrease task-induced BOLD signal changes in the ATL. Second, the effects of cTBS on semantic processing can be attributed to GABAergic action in the ATL at the individual level: greater changes in GABA concentrations would result in increasing changes in task-induced BOLD signal during semantic processing and semantic task performance. Furthermore, to address and explore the relationship between regional GABA levels in the ATL and semantic memory function, we combined data from our previous study (47) with the current study’s data. We then explored the function of GABA in the ATL in relation to semantic function. Finally, we extended this GABAergic function to semantic neuroplasticity driven by cTBS.

Results

We acquired resting-state MRS for the left ATL and vertex followed by fMRI before and after cTBS (Fig. 1A). During fMRI, participants made semantic association decisions as an active task and pattern matching as a control task (Fig. 1B). GABA concentrations were estimated from the ATL with the vertex as a control region (Fig. 1C). cTBS with 80% of resting motor threshold (RMT) was delivered outside of scanner at one of the target regions with a week gap between two sessions (Fig. 1D).

A) Experimental procedure. B) An example of the semantic association task (left) and control task (right: pattern matching). Each trial starts with a fixation followed by stimuli, which have 3 items, a target on the top and 2 choices at the bottom. C) The location of volume of interest (VOI) for MRS (left ATL and vertex) and a representative MRS spectrum with estimated peaks (right). Colour bar indicates the number of overlapping participants. NAA: N-acetylaspartate. D) cTBS protocols. cTBS was applied over the left ATL and vertex as a control site. Each stimulation was delivered on different days with a week gap at least.

cTBS modulates regional GABA concentrations and task-related BOLD signal changes in the ATL

E-field modelling of cTBS showed that, as intended, ATL cTBS stimulated the left ventrolateral ATL (Fig. 2A). To investigate how cTBS modulates GABA concentrations, we quantified GABA/NAA and calculated the changes (POST– PRE). A 2 × 2 repeated measures analysis of variance (ANOVA) with stimulation (ATL vs. vertex) and VOI (ATL vs. vertex) as within subject factor was performed. There was a significant interaction effect between the stimulation and VOI (F1,16 = 4.57, p = 0.048) (Fig. 2B). There was no significant main effect of the stimulation (F1,16 = 3.23, p = 0.091) and VOI (F1,16 = 0.64, p = 0.435). Planned paired t-tests revealed that ATL stimulation significantly increased GABA concentrations in the ATL compared to the control stimulation (t = 1.86, p =0.040) and control site (t = 2.07, p = 0.027). There were no cTBS effects in the vertex VOI regardless of the stimulation (ps > 0.23). It is noted that there was no significant cTBS effect in Glx (Fig. S1).

Removed in the revised version: The effects of cTBS in the ATL.

Removed in the revised version: A) ATL cTBS e-field modelling. B) cTBS-induced regional GABA changes in the ATL. Red bar indicates the ATL stimulation and white bar indicates the control (vertex) stimulation. C) fMRI results of the contrast of interest (semantic > control) in the ATL pre-stimulation session. D) cTBS-induced ATL BOLD signal changes during a semantic and control task. White bars represent the pre-stimulation session, and grey bars represent the post-stimulation session. E) The relationship between cTBS-induced GABA changes and BOLD signal changes in the ATL. F) The results of task performance. A positive value of cTBS effect (Post – Pre) in IE suggests an inhibitory effect, indicating poorer performance after the stimulation. In contrast, a negative value denotes a facilitatory effect, signifying improved performance following the stimulation. Red bar indicates the ATL stimulation and white bar indicates the control (vertex) stimulation. Each individual is represented as a circle. * p < 0.05, ** p < 0.01

fMRI results demonstrated that the semantic association task evoked increased activation in the ATL, prefrontal and posterior temporal cortex compared to the control task (Fig. 2C, Fig. S2 and Table S1). To examine the effects of cTBS, we performed ROI analysis using the same VOI in the ATL. Planned paired t-tests revealed that BOLD signal changes during semantic processing were significantly altered after ATL cTBS compared to the pre-stimulation (t = 1.78, p = 0.046) and the control stimulation (t = -2.11, p = 0.025) (Fig. 2D).

To investigate the effects of ATL cTBS, we conducted a partial correlation analysis between GABA changes (POST–PRE) and BOLD signal changes (POST– PRE), accounting for age and sex. We found a significant correlation between cTBS induced GABA changes and BOLD signal changes in the ATL (r = -0.86, p < 0.001). Individuals with greater increases in ATL GABA levels following ATL cTBS showed greater reductions in task-induced BOLD signal changes in the ATL. These results demonstrate that ATL cTBS modifies regional GABA concentrations, and the cTBS-induced changes in GABA levels are connected to individual-level changes in task-related fMRI signal.

To evaluate the cTBS effects in behaviour, we used inverse efficiency (IE) score (RT/1-the proportion of error) and calculated IE changes (POST-PRE). The cTBS effects on participants’ performance was examined using a 2 × 2 repeated measures ANOVA with stimulation (ATL vs. vertex) and task (semantic vs. control) as within-subject factors. There was no significant main effect of stimulation. However, we found a significant main effect of task (F1, 15 = 6.66, p = 0.021) and a marginally significant interaction between the stimulation and task (F1, 15 = 4.06, p = 0.061). Post hoc paired t-tests demonstrated that participants performed semantic task worse (higher IE score) after the ATL stimulation compared to the control task (t = 2.81, p = 0.006) and vertex stimulation (t = 1.91, p = 0.038) (Fig. 2F). The results showed that ATL cTBS induced the task-specific inhibitory effects on semantic task performance. It is noted that higher IE score indicates poorer performance. The results of accuracy and RT for each task were summarised in the Supplementary Table 2 and Figure S3.

Added in the revised version: The effects of cTBS in the ATL.

Added in the revised version: A) ATL cTBS e-field modelling. B) cTBS-induced regional GABA changes in the ATL. Red bar indicates the ATL stimulation and white bar indicates the control (vertex) stimulation. C) fMRI results of the contrast of interest (semantic > control) in the ATL pre-stimulation session. D) cTBS-induced ATL BOLD signal changes during a semantic and control task. White bars represent the pre-stimulation session, and grey bars represent the post-stimulation session. E) The relationship between cTBS-induced GABA changes and BOLD signal changes in the ATL. F) The results of task performance. A positive value of cTBS effect (Post – Pre) in IE suggests an inhibitory effect, indicating poorer performance after the stimulation. In contrast, a negative value denotes a facilitatory effect, signifying improved performance following the stimulation. Red bar indicates the ATL stimulation and white bar indicates the control (vertex) stimulation. Each individual is represented as a circle. * p < 0.05, ** p < 0.01

Earlier version: Moreover, we categorised participants based on changes in their semantic task performance following ATL stimulation to examine the relationship between pre-stimulation ATL GABA levels and cTBS-induced behavioural changes. After ATL stimulation, responders exhibited poorer semantic task performance (higher IE) (t = - 1.937, p = 0.050), whereas non-responders demonstrated a paradoxical, facilitatory effects on semantic task performance (lower IE) (t = 2.872, p = 0.009) (Fig. 3A). Notably, both responders and non-responders showed increased GABA levels in the ATL following stimulation (responder: t = -2.203, p = 0.035, non-responder: t = -3.912, p = 0.001) (Fig. 3B). Our planned t-tests revealed that there was a significant difference between responders and non-responders in the semantic task performance (t = -1.76, p = 0.050) and ATL GABA levels in pre-stimulation session (t = 2.78, p = 0.006) (Fig. 3). Responders showed better semantic task performance with higher ATL GABA levels compared to non-responders.

Revised version: Moreover, we categorised participants based on changes in their semantic task performance following ATL stimulation to examine the relationship between ATL GABA levels and individual behavioural response to cTBS. To account for practice effects, we first adjusted task performance (IE) by normalizing ATL stimulation performance relative to with vertex stimulation performance. A 2×2×2 ANOVA was conducted with task (semantic vs. control) and session (PRE vs. POST) as within-subject factors, and group (responders vs. non-responders) as a between-subject factor. The analysis revealed a significant interaction between the session and group (F1, 15 = 10.367, p = 0.006), a marginally significant interaction between the session and task (F1, 15 = 4.370, p = 0.054), and a significant three-way interaction between the session, task, and group (F1, 15 = 7.580, p = 0.015). Post hoc tests showed that after ATL stimulation, responders exhibited poorer semantic task performance (higher IE) (t = -5.281, p < 0.001), whereas non-responders demonstrated a paradoxical, facilitatory effects on semantic task performance (lower IE) (t = 3.206, p = 0.007) (Fig. 3A). Additionally, responders showed pooper semantic task performance compared to non-responders after ATL stimulation (t = 2.349, p = 0.033) (Fig. 3A). Notably, no differences were observed between responders and non-responders in the control task performance across pre- and post-stimulation sessions, confirming that the practice effect was successfully controlled (Fig. 3B). A 2 x 2 ANOVA with session (PRE vs. POST) as a within subject factor and with group (responders vs. non-responders) as a between subject factor was conducted to investigate the effects of individual cTBS responsiveness on ATL GABA levels. The analysis revealed a significant main effect of session (F1, 14 = 39.906, p < 0.001) and group (F1, 14 = 9.677, p = 0.008). Post hoc paired t-tests revealed that both responders and non-responders showed increased GABA levels in the ATL following stimulation (responder: t = -3.885, p = 0.002, non-responder: t = -4.831, p = 0.001) (Fig. 3C). Post hoc t-tests further revealed a significant difference in ATL GABA levels between responders and non-responders in both pre-stimulation (t = 2.816, p = 0.007) and post-stimulation session (t = 2.555, p = 0.011) (Fig. 3C).

Earlier version: A) Semantic task performance in pre- and post-ATL stimulation session. B) ATL GABA levels in pre- and post-ATL stimulation session. The red circle represents the responder, while the blue diamond denotes the non-responder. Each individual is represented as a circle. Error bars indicates standard errors. * p ≤ 0.05, ** p < 0.01, *** p < 0.001

Revised version: A) Semantic task performance (IE) in pre- and post-ATL stimulation session. B) Control task performance (IE) in pre-and post- ATL stimulation session. To account for the practice effect, task performance was adjusted by dividing ATL stimulation performance by vertex stimulation performance. C) ATL GABA levels in pre- and post-ATL stimulation session. The red circle represents the responder, while the blue diamond denotes the non-responder. Each individual is represented as a circle. Error bars indicate standard errors. * p < 0.05, ** p < 0.01, *** p < 0.001

Regional GABA concentrations in the ATL play a crucial role in semantic memory

Earlier version: In our prior study (48), ATL GABA levels were significantly and negatively correlated with ATL activity during semantic processing (Fig. 4A). Here, we replicated our previous findings in a different cohort with the same research paradigm (pre-stimulation session). We conducted a single-voxel regression analysis with the individual’s GABA concentrations (ATL pre-stimulation session) as the regressor of the fMRI contrast of interest (semantic > control). The BOLD response in the ventral ATL was significantly and negatively correlated with the individual GABA levels in the ATL (MNI -42 -6 -33, p SVC-FWE < 0.05), overlapping with the results from our previous study (48) (Fig. 4A). The GABA-related region of the ventral ATL overlapped with the semantic coding hotspot from electrocorticograms (ECoG) data and direct cortical stimulation (31, 32) (Fig. 4A). Furthermore, we found that individual GABA concentrations in the ATL were positively associated with semantic task performance (Fig. 4B). We also confirmed this finding, demonstrating that individuals with more GABA in the ATL performed the semantic task better (higher accuracy) (r = 0.50, p = 0.035) (Fig. 4B). It should be noted that individual GABA levels also significantly correlated with ATL activity and semantic task performance at the vertex stimulation session (Fig. S4). There was no significant relationship between ATL GABA levels and RT during semantic processing (ps > 0.44) (Fig. S5) and between ATL GABA levels and control task performance (Table. S3). These results demonstrate that higher levels of cortical GABA in the ATL are associated with task-related regional activity as well as enhanced semantic function.

Revised version: In our prior study (47), ATL GABA levels were significantly and negatively correlated with ATL activity during semantic processing (Fig. 4A). Here, we replicated our previous findings in a different cohort with the same research paradigm (pre-stimulation session). We conducted a single-voxel regression analysis with the individual’s GABA concentrations (ATL pre-stimulation session) as the regressor of the fMRI contrast of interest (semantic > control). The BOLD response in the ventral ATL was significantly and negatively correlated with the individual GABA levels in the ATL (MNI -42 -6 -33, p SVC-FWE < 0.05), overlapping with the results from our previous study (47) (Fig. 4A). The GABA-related region of the ventral ATL overlapped with the semantic coding hotspot from electrocorticograms (ECoG) data and direct cortical stimulation (30, 31) (Fig. 4A). Furthermore, we found that individual GABA concentrations in the ATL were positively associated with semantic task performance (Fig. 4B). We also confirmed this finding, demonstrating that individuals with more GABA in the ATL performed the semantic task better (higher accuracy) (r = 0.50, p = 0.035) (Fig. 4B). It should be noted that individual GABA levels also significantly correlated with ATL activity and semantic task performance at the vertex stimulation session (Fig. S4). There was no significant relationship between ATL GABA levels and RT during semantic processing (ps > 0.44) (Fig. S5) and between ATL GABA levels and control task performance (Table. S3). These results demonstrate that higher levels of cortical GABA in the ATL are associated with task-related regional activity as well as enhanced semantic function.

A) Local maxima of the voxel-wise regression analysis of the contrast (semantic > control) with GABA concentrations in the ATL. B) The relationship between individual GABA levels in the ATL and semantic task performance from our previous study (Jung et al., 2017) and current study (pre-stimulation session). C) The ATL GABA function in relation to semantic performance. D) The relationship between cTBS-induced changes in ATL GABA levels and semantic task performance. Dotted line represents the linear function between ATL GABA levels and semantic task performance. Coloured line represents the inverted U-shaped (quadratic) function between ATL GABA levels and semantic task performance.

The inverted U-shaped function of ATL GABA concentrations in semantic processing

Earlier version: The pattern of correlation between GABAergic activity and semantic task accuracy observed in our previous study was replicated in an entirely new cohort in the current study. Next, we combined the two studies (N = 37) in order to fully investigate the potential role of GABA in the ATL as a mechanistic link between ATL inhibitory GABAergic action and semantic task performance (accuracy). First, we tested the linear relationship between ATL GABA levels and semantic task performance. We confirmed our previous findings that individuals with higher GABA levels in the ATL showed better semantic task performance (R2 = 0.49, p < 0.001) (Fig. 4C). Second, to test our hypothesis, we assumed that semantic performance follows an inverted U-shaped (quadratic) function with relation to ATL GABA concentrations. In other words, people who have low or excessive GABA levels in the ATL perform the semantic task relatively poorly. The results revealed that the inverted U-shaped function between ATL GABA and semantic performance was significant (R2 = 0.67, p < 0.001) (Fig. 4C). To compare two different models, we calculated the Bayesian Information Criterion (BIC) as a measure of model fitness (51) and performed a partial F-test to determine whether there is a statistically significant difference between two models. A best model fitness can be characterized by low BIC and high R2. The results showed a BIC value of 243.72 for the linear function and a value of 233.36 for the quadratic function. The results of F-tests revealed that the inverted U-shaped model provided a statistically significantly better fit than the linear model (F = 15.60, p < 0.001). The best-fitting model is therefore the inverted-U-shaped function of ATL GABA in semantic processing. There was no significant relationship between ATL GABA levels and RT during semantic processing (linear function R2 = 0.21, p =0.45, quadratic function: R2 = 0.17, p = 0.21).

Revised version: The pattern of correlation between GABAergic activity and semantic task accuracy observed in our previous study was replicated in an entirely new cohort in the current study. Next, we combined the two studies (N = 37) in order to fully investigate the potential role of GABA in the ATL as a mechanistic link between ATL inhibitory GABAergic action and semantic task performance (accuracy). First, we tested the linear relationship between ATL GABA levels and semantic task performance. We confirmed our previous findings that individuals with higher GABA levels in the ATL showed better semantic task performance (R2 = 0.49, p < 0.001) (Fig. 4C). Second, to test our hypothesis, we assumed that semantic performance follows an inverted U-shaped (quadratic) function with relation to ATL GABA concentrations. In other words, people who have low or excessive GABA levels in the ATL perform the semantic task relatively poorly. The results revealed that the inverted U-shaped function between ATL GABA and semantic performance was significant (R2 = 0.67, p < 0.001) (Fig. 4C). To compare two different models, we calculated the Bayesian Information Criterion (BIC) as a measure of model fitness (50) and performed a partial F-test to determine whether there is a statistically significant difference between two models. A best model fitness can be characterized by low BIC and high R2. The results showed a BIC value of 243.72 for the linear function and a value of 233.36 for the quadratic function. The results of F-tests revealed that the inverted U-shaped model provided a statistically significantly better fit than the linear model (F = 15.60, p < 0.001). The best-fitting model is therefore the inverted-U-shaped function of ATL GABA in semantic processing. There was no significant relationship between ATL GABA levels and RT during semantic processing (linear function R2 = 0.21, p =0.45, quadratic function: R2 = 0.17, p = 0.21).

We performed the same analysis on the pre- and post-stimulation data in order to investigate the role of ATL GABA in semantic plasticity. We found that there was a significant linear relationship between the ATL GABA levels and semantic performance before and after stimulation (R2 = 0.19, p < 0.01) (Fig. 4D). The inverted U-shaped function also showed a significant association between them (R2 = 0.33, p < 0.005) (Fig. 4D). The F-test demonstrated that the quadratic model showed a significantly better fit than the linear model (F = 6.64, p = 0.014). The inverted U-shaped function has the better BIC score for explaining changes in ATL GABA levels and semantic performance induced by cTBS (linear model BIC 230.21, quadratic model BIC 227.13). Thus, the best-fitting model is the inverted U-shaped for the ATL GABA changes induced by cTBS in relation to semantic function.

Discussion

Earlier version: We investigated the role of cortical GABA in the ATL on semantic memory and its neuroplasticity. Our results demonstrated an increase in regional GABA levels following inhibitory cTBS in human associative cortex, specifically in the ATL, a representational semantic hub. Notably, the observed increase was specific to the ATL and semantic processing, as it was not observed in the control region (vertex) and not associated with control processing (visuospatial processing). Our study also found that the magnitude of cTBS-modulated GABA changes at the individual level was associated with their changes in ATL activity during semantic processing. Furthermore, our data confirmed and replicated our previous findings that GABA concentrations in the ATL shape task-related cortical activity and semantic task performance. In other words, individuals with greater semantic performance exhibit selective activity in the ATL due to higher concentrations of inhibitory GABA. GABAergic inhibition can sharpen activated distributed semantic representations through lateral inhibition, leading to improved semantic acuity (48), which aligns with theories on representational sharpening in visual perception (52, 53). Importantly, our data revealed, for the first time, a non-linear, inverted-U-shape relationship between GABA levels in the ATL and semantic function, by explaining individual differences in semantic task performance and cTBS responsiveness. Understanding the link between neurochemistry and semantic memory is an important step in understanding individual differences in semantic behaviour and could guide therapeutic interventions to restore semantic abilities in clinical settings.

Revised version: We investigated the role of cortical GABA in the ATL on semantic memory and its neuroplasticity. Our results demonstrated an increase in regional GABA levels following inhibitory cTBS in human associative cortex, specifically in the ATL, a representational semantic hub. Notably, the observed increase was specific to the ATL and semantic processing, as it was not observed in the control region (vertex) and not associated with control processing (visuospatial processing). Our study also found that the magnitude of cTBS-modulated GABA changes at the individual level was associated with their changes in ATL activity during semantic processing. Furthermore, our data confirmed and replicated our previous findings that GABA concentrations in the ATL shape task-related cortical activity and semantic task performance. In other words, individuals with greater semantic performance exhibit selective activity in the ATL due to higher concentrations of inhibitory GABA. GABAergic inhibition could sharpen activated distributed semantic representations through lateral inhibition, leading to improved semantic acuity (47), which aligns with theories on representational sharpening in visual perception (51, 52). Importantly, our data revealed, for the first time, a non-linear, inverted-U-shape relationship between GABA levels in the ATL and semantic function, by explaining individual differences in semantic task performance and cTBS responsiveness. Understanding the link between neurochemistry and semantic memory is an important step in understanding individual differences in semantic behaviour and could guide therapeutic interventions to restore semantic abilities in clinical settings.

Earlier version: To the best of our knowledge, this is the first study to demonstrate that (1) cTBS modulates both regional GABA concentrations and cortical activity in human higher cognition - semantic memory, and that (2) changes in GABA levels are closely linked to changes in regional activity induced by cTBS. These results suggests that GABAergic activity may be the mechanism by which cTBS induces long-lasting after-effects on cortical excitability, leading to behavioural changes. Previous studies in animals and humans have also suggested that cTBS can induce LTD-like effects on cortical synapses and is associated with the GABAergic system in the cortex (54-59). Another study employing MRS found that cTBS increased regional GABA concentrations at the primary motor cortex in healthy subjects (60). These findings suggest that cTBS activates a population of cortical GABAergic interneurons, leading to the increase in GABAergic activity (61, 62). As a major inhibitory neurotransmitter, GABA has been shown to have a negative correlation with BOLD signal changes (5, 63). Previously we demonstrated this negative relationship between ATL GABA levels and BOLD signal changes in the ATL during semantic processing (48), indicating a potential role of GABA in shaping the functions/computations of the cortex. Here, we further demonstrated that increase in GABA induced by cTBS was negatively correlated with the reduction of BOLD signal responses in the ATL following cTBS, during semantic processing. Our findings suggest a crucial role for GABAergic inhibition in the ATL shaping the local neural functioning underpinning semantic memory and its neuroplasticity. The GABAergic inhibition confines the propagation of excitatory signalling, thereby maintaining the functional organization of the cortex (64), and the modulation of cortical GABAergic inhibition drives experience-dependent plasticity in cognition (10, 65).

Revised version: To the best of our knowledge, this is the first study to demonstrate that (1) cTBS modulates both regional GABA concentrations and cortical activity in human higher cognition - semantic memory, and that (2) changes in GABA levels are closely linked to changes in regional activity induced by cTBS. These results suggests that GABAergic activity may be the mechanism by which cTBS induces long-lasting after-effects on cortical excitability, leading to behavioural changes. Previous studies in animals and humans have also suggested that cTBS can induce LTD-like effects on cortical synapses and is associated with the GABAergic system in the cortex (53-58). Another study employing MRS found that cTBS increased regional GABA concentrations at the primary motor cortex in healthy subjects (59). These findings suggest that cTBS activates a population of cortical GABAergic interneurons, leading to the increase in GABAergic activity (60, 61). As a major inhibitory neurotransmitter, GABA has been shown to have a negative correlation with BOLD signal changes (5, 62). Previously we demonstrated this negative relationship between ATL GABA levels and BOLD signal changes in the ATL during semantic processing (47), indicating a potential role of GABA in shaping the functions/computations of the cortex. Here, we further demonstrated that increase in GABA induced by cTBS was negatively correlated with the reduction of BOLD signal responses in the ATL following cTBS, during semantic processing. Our findings suggest a crucial role for GABAergic inhibition in the ATL shaping the local neural functioning underpinning semantic memory and its neuroplasticity. The GABAergic inhibition confines the propagation of excitatory signalling, thereby maintaining the functional organization of the cortex (63), and the modulation of cortical GABAergic inhibition drives experience-dependent plasticity in cognition (10, 64).

Earlier version: GABA exists in two distinct neuronal pools: cytoplasmic GABA, which is involved in metabolism, and vesicular GABA, which plays a role in inhibitory synaptic neurotransmission (66). In addition to intracellular GABA, extracelluar GABA exerts tonic inhibition through extra-synaptic GABAA receptors (67). MRS is capable of detecting the total concentration of GABA in the voxel of interest, but it cannot differentiate between different pools of GABA (68, 69). Some studies have suggested that MRS-measured GABA signals reflect GABAergic tonic inhibition rather than synaptic GABA signalling (70, 71) whereas other studies have failed to replicate this relationship (72, 73). A recent study has shown a link between MRS-measured GABA and phasic synaptic GABAergic activity (74). Although findings of previous studies have been mixed, changes in GABA levels observed in this study may reflect cTBS-modulated GABAergic neurotransmission, which encompass both tonic and synaptic GABAergic activity. This GABAergic activity shapes the selective response profiles of neurons in the cortex (9).

Revised version: GABA exists in two distinct neuronal pools: cytoplasmic GABA, which is involved in metabolism, and vesicular GABA, which plays a role in inhibitory synaptic neurotransmission (65). In addition to intracellular GABA, extracelluar GABA exerts tonic inhibition through extra-synaptic GABAA receptors (66). MRS is capable of detecting the total concentration of GABA in the voxel of interest, but it cannot differentiate between different pools of GABA (67, 68). Some studies have suggested that MRS-measured GABA signals reflect GABAergic tonic inhibition rather than synaptic GABA signalling (69, 70) whereas other studies have failed to replicate this relationship (71, 72). A recent study has shown a link between MRS-measured GABA and phasic synaptic GABAergic activity (73). Although findings of previous studies have been mixed, changes in GABA levels observed in this study may reflect cTBS-modulated GABAergic neurotransmission, which encompass both tonic and synaptic GABAergic activity. This GABAergic activity shapes the selective response profiles of neurons in the cortex (9).

Earlier version: Inverted U-shaped models have been previously considered in the field of neuroscience, specifically in terms of the relationship between the concentration of neurotransmitters such as dopamine, acetylcholine and noradrenaline, and the level of neural activity (75-78). Recent studies suggest that this relationship also applies to behaviour, where moderate levels of neural activity are linked to the optimal performance (for a reveiw, see 79). For example, Ferri et al. (80) showed an inverted U-shaped relationship between excitation and inhibition balance and multisensory integration, where extreme values impair functionality while intermediate values enhance it, even in healthy individuals. Our findings revealed a non-linear relationship between GABA levels in the anterior temporal lobe and semantic function, indicating that individual variations in semantic task performance can be explained by an inverted-U-shape pattern (Fig. 5A). Specifically, for relatively greater levels of GABA in the ATL, with lower task-induced regional activity, were associated with better semantic processing in healthy participants (48). That is, individuals with better semantic memory abilities show more specific cortical activity in the ATL, which is linked to higher concentrations of inhibitory GABA. Extreme levels of GABA can be found in studies with dementia patients and pharmacological studies with GABA agonists. Recent studies have reported decreased GABA levels in Alzheimer’s disease (81, 82) and frontotemporal dementia (83, 84) in relation to their cognitive impairments such as memory and language. In fact, GABA agonists like midazolam have been found to improve a verbal generation in anxiety patients by increasing GABAergic function (85). On the other hand, healthy participants who received GABA supplementation (such as baclofen) have been found to have decreased task performance (86). Overall, optimal, elevated levels of GABA in the ATL may aid in refining stimulated widespread semantic representations through local inhibitory processes.

Revised version: Inverted U-shaped models have been previously considered in the field of neuroscience, specifically in terms of the relationship between the concentration of neurotransmitters such as dopamine, acetylcholine and noradrenaline, and the level of neural activity (74-77). Recent studies suggest that this relationship also applies to behaviour, where moderate levels of neural activity are linked to the optimal performance (for a reveiw, see 78). For example, Ferri et al. (79) showed an inverted U-shaped relationship between excitation and inhibition balance and multisensory integration, where extreme values impair functionality while intermediate values enhance it, even in healthy individuals. Our findings revealed a non-linear relationship between GABA levels in the anterior temporal lobe and semantic function, indicating that individual variations in semantic task performance can be explained by an inverted-U-shape pattern (Fig. 5A). Specifically, for relatively greater levels of GABA in the ATL, with lower task-induced regional activity, were associated with better semantic processing in healthy participants (47). That is, individuals with better semantic memory abilities show more specific cortical activity in the ATL, which is linked to higher concentrations of inhibitory GABA. Extreme levels of GABA can be found in studies with dementia patients and pharmacological studies with GABA agonists. Recent studies have reported decreased GABA levels in Alzheimer’s disease (80, 81) and frontotemporal dementia (82, 83) in relation to their cognitive impairments such as memory and language. In fact, GABA agonists like midazolam have been found to improve a verbal generation in anxiety patients by increasing GABAergic function (84). On the other hand, healthy participants who received GABA supplementation (such as baclofen) have been found to have decreased task performance (85). Overall, optimal, elevated levels of GABA in the ATL may aid in refining stimulated widespread semantic representations through local inhibitory processes. However, our findings should be interpreted with caution due to the limitation of having fewer data points in the latter half (right side) of the inverted U-shaped curve. Future studies incorporating GABA agonists could help further validate and refine these findings

Schematic diagram of relationship between the concentration of GABA in the ATL and semantic function.

A) Inverted U-shaped ATL GABA function in semantic memory B) Inverted U-shaped ATL GABA function for cTBS response on semantic memory.

Earlier version: This inverted U-shaped model could also explain inter-individual variability in cTBS-induced neuroplasticity in the ATL in semantic processing. Our data demonstrated that cTBS over ATL increased regional GABA concentrations, but there was inter-individual variability in GABA level changes in response to cTBS (Fig. 2). Our previous investigation (8) showed that the pre-interventional neurochemical state was crucial in predicting cTBS-induced changes in semantic memory. Specifically, cTBS over the ATL inhibited the semantic task performance (i.e., reduced accuracy) of individuals with initially higher concentration of GABA in the ATL, linked to better semantic capacity. However, cTBS had a null or even facilitatory effect on individuals with lower semantic ability with relatively lower GABA levels in the ATL. This study suggests that individuals with higher GABA levels in the ATL were more likely to respond to cTBS, exhibiting inhibitory effects on semantic task performance (responders), while individuals with lower GABA concentrations and lower semantic ability were less likely to respond or even showed facilitatory effects after ATL cTBS (non-responders). The current study revealed a non-linear, inverted-U-shape relationship between GABA levels in the ATL and semantic function, by explaining individual differences in semantic task performance and cTBS responsiveness. As regional GABA increases after cTBS, responders with the optimal level of GABA in the ATL would show poorer semantic performance, whereas non-responders could exhibit no changes or even better semantic performance with GABA increase (Fig. 5B). This relationship is similar to the inverted U-shaped relationship between dopamine action in the prefrontal cortex (PFC) and cognitive control, whereby moderate levels of dopamine lead to optimal cognitive performance (87). The effects of dopaminergic drug on PFC function also depend on baseline levels of working memory performance (for a review, see 76), explaining the effects of dopaminergic drugs on cognitive performance in individuals with varying working memory capacities (88-90).

Revised version: This inverted U-shaped model could also explain inter-individual variability in cTBS-induced neuroplasticity in the ATL in semantic processing. Our data demonstrated that cTBS over ATL increased regional GABA concentrations, but there was inter-individual variability in GABA level changes in response to cTBS (Fig. 2). Our previous investigation (8) showed that the pre-interventional neurochemical state was crucial in predicting cTBS-induced changes in semantic memory. Specifically, cTBS over the ATL inhibited the semantic task performance (i.e., reduced accuracy) of individuals with initially higher concentration of GABA in the ATL, linked to better semantic capacity. However, cTBS had a null or even facilitatory effect on individuals with lower semantic ability with relatively lower GABA levels in the ATL. This study suggests that individuals with higher GABA levels in the ATL were more likely to respond to cTBS, exhibiting inhibitory effects on semantic task performance (responders), while individuals with lower GABA concentrations and lower semantic ability were less likely to respond or even showed facilitatory effects after ATL cTBS (non-responders). The current study revealed a non-linear, inverted-U-shape relationship between GABA levels in the ATL and semantic function, by explaining individual differences in semantic task performance and cTBS responsiveness. As regional GABA increases after cTBS, responders with the optimal level of GABA in the ATL would show poorer semantic performance, whereas non-responders could exhibit no changes or even better semantic performance with GABA increase (Fig. 5B). This relationship is similar to the inverted U-shaped relationship between dopamine action in the prefrontal cortex (PFC) and cognitive control, whereby moderate levels of dopamine lead to optimal cognitive performance (86). The effects of dopaminergic drug on PFC function also depend on baseline levels of working memory performance (for a review, see 75), explaining the effects of dopaminergic drugs on cognitive performance in individuals with varying working memory capacities (87-89).

Our findings provide novel evidence of a direct link from neurochemical modulations to cortical responses in the brain, highlighting substantial individual variability in semantic memory and plasticity. In addition, the current study represents an important replication and extension of previous findings regarding the role of GABAergic inhibition in semantic memory. These results offer fundamental insights into the mechanisms underlying the maintenance and alteration of functional cortical organization in response to perturbations. Our study has important implications for the development of personalized therapeutic interventions aimed at modulating neurochemical systems to restore or enhance higher cognitive function in humans.

Methods and materials

Participants

Earlier version: Nineteen healthy English native speakers (9 females, mean age = 25.9 ± 5.8 years, age range: 19-38) participated in this study. The sample size was calculated based on a previous study (39), which indicated that to achieve α=0.05, power=80% for the critical interaction between TMS and task then N≥17 were required. A participant completed one session (ATL stimulation) only. All participants were right-handed (91). All subjects provided informed written consent. The study was approved by the local ethics committee.

Revised version: Nineteen healthy English native speakers (9 females, mean age = 25.9 ± 5.8 years, age range: 19-38) participated in this study. The sample size was calculated based on a previous study (38), which indicated that to achieve α=0.05, power=80% for the critical interaction between TMS and task then N≥17 were required. A participant completed one session (ATL stimulation) only. All participants were right-handed (90). All subjects provided informed written consent. The study was approved by the local ethics committee.

Earlier version: To explore the role of GABA in semantic memory function, we used the data previously published (48). Data from twenty healthy, right-handed native English speakers were included (7 males, mean age = 23 ± 4 years, age range: 20-36).

Revised version: To explore the role of GABA in semantic memory function, we used the data previously published (47). Data from twenty healthy, right-handed native English speakers were included (7 males, mean age = 23 ± 4 years, age range: 20- 36).

Experimental design and procedure

Participants were asked to visit two times for the study. In each visit, the target region was identified prior to the baseline scan. Participants had a multimodal imaging (MRS and fMRI). Then, participants were removed from the scanner and cTBS was performed in a separate room. Following cTBS, participants were repositioned into the scanner and had the second multimodal imaging (Fig. 1A).

Earlier version: We used the same paradigm for the multimodal imaging from our previous study (48). During MRS, participants were asked to be relaxed with eyes open. Participants performed a semantic association decision task and pattern matching as a control task during fMRI scanning (Fig. 1B). The semantic association decision task required a participant to choose which of two pictures at the bottom of the screen was more related in meaning to a probe picture presented on the top of the screen. The items for the semantic association task were from the Pyramids and Palm Tree test (92) and Camel and Cactus test (28). Items for the pattern matching task were created by scrambling the pictures used in the semantic association task. In the pattern matching task, a participant was asked to identify which of two patterns at the bottom was visually identical to a probe pattern on the top (Fig. 1B). Participants were required to press one of two buttons designating two choices in a trial. In each trial, there was a fixation for 500ms followed by the stimuli for 4500ms. A task block had four trials of each task. There were 9 blocks of each task interleaved (e.g., A-B-A-B) with a fixation for 4000ms during fMRI. Total scanning time was about 8mins. E-prime software (Psychology Software Tools Inc., Pittsburgh, USA) was used to display stimuli and to record responses.

Revised version: We used the same paradigm for the multimodal imaging from our previous study (47). During MRS, participants were asked to be relaxed with eyes open. Participants performed a semantic association decision task and pattern matching as a control task during fMRI scanning (Fig. 1B). The semantic association decision task required a participant to choose which of two pictures at the bottom of the screen was more related in meaning to a probe picture presented on the top of the screen. The items for the semantic association task were from the Pyramids and Palm Tree test (91) and Camel and Cactus test (27). Items for the pattern matching task were created by scrambling the pictures used in the semantic association task. In the pattern matching task, a participant was asked to identify which of two patterns at the bottom was visually identical to a probe pattern on the top (Fig. 1B). Participants were required to press one of two buttons designating two choices in a trial. In each trial, there was a fixation for 500ms followed by the stimuli for 4500ms. A task block had four trials of each task. There were 9 blocks of each task interleaved (e.g., A-B-A-B) with a fixation for 4000ms during fMRI. Total scanning time was about 8mins. E-prime software (Psychology Software Tools Inc., Pittsburgh, USA) was used to display stimuli and to record responses.

Transcranial magnetic stimulation

Earlier version: A Magstim Super Rapid stimulator (MagStim Company, Whitland, UK) with a figure of eight coil (70mm standard coil) was used to deliver cTBS over the left ATL or vertex with a week gap between the stimulation (Fig. 1D). cTBS consisted of bursts containing 3 pulses at 50Hz (50) and was applied at 80% of the resting motor threshold (RMT), previously showed inhibitory effects on semantic processing in the ATL (39). RMT was established for each individual, defined as the minimum intensity of stimulation required to produce twitches on 5 of 10 trials from the right first dorsal interosseous (FDI) muscle when the participant was at rest. The average stimulation intensity (80% RMT) was 49.2% ranging from 38% to 60%. The target site in the left ATL was delineated based on the peak coordinate (MNI -36 -15 -30), which represents maximal peak activation observed during semantic processing in previous distortion-corrected fMRI studies (38, 41). This coordinate was transformed to each individual’s native space using Statistical Parametric Mapping software (SPM8, Wellcome Trust Centre for Neuroimaging, London, UK). T1 images were normalised to the MNI template and then the resulting transformations were inverted to convert the target MNI coordinate back to the individual’s untransformed native space coordinate. These native-space ATL coordinates were subsequently utilized for frameless stereotaxy, employing the Brainsight TMS-MRI co-registration system (Rogue Research, Montreal, Canada). The vertex (Cz) was designated as a control site following the international 10–20 system.

Revised version: A Magstim Super Rapid stimulator (MagStim Company, Whitland, UK) with a figure of eight coil (70mm standard coil) was used to deliver cTBS over the left ATL or vertex with a week gap between the stimulation (Fig. 1D). cTBS consisted of bursts containing 3 pulses at 50Hz (49) and was applied at 80% of the resting motor threshold (RMT), previously showed inhibitory effects on semantic processing in the ATL (38). RMT was established for each individual, defined as the minimum intensity of stimulation required to produce twitches on 5 of 10 trials from the right first dorsal interosseous (FDI) muscle when the participant was at rest. The average stimulation intensity (80% RMT) was 49.2% ranging from 38% to 60%.

Added in the revised version: Previous rTMS studies targeted a lateral ATL site 10 mm posterior to the temporal pole on the middle temporal gyrus (MTG) (41, 92, 93), aligning with the broader ATL region typically atrophied in semantic dementia (40). However, distortion-corrected fMRI studies (37, 40) have revealed graded activation differences across the ATL, with peak activation in the ventromedial ATL. Based on these findings, we selected the target site in the left ATL (MNI -36 -15 -30) from a prior distortion-corrected fMRI study (37, 40) that employed the same tasks as our study (for further details, see the Supplementary Information). This coordinate was transformed to each individual’s native space using Statistical Parametric Mapping software (SPM8, Wellcome Trust Centre for Neuroimaging, London, UK). T1 images were normalised to the MNI template and then the resulting transformations were inverted to convert the target MNI coordinate back to the individual’s untransformed native space coordinate. These native-space ATL coordinates were subsequently utilized for frameless stereotaxy, employing the Brainsight TMS-MRI co-registration system (Rogue Research, Montreal, Canada). The vertex (Cz) was designated as a control site following the international 10–20 system.

Earlier version: SimNIBS 3.2 (93) was used to calculate individual electric field of cTBS. The pipeline by Nielsen et al (94) was utilized to generate the individual head model consisting of five tissue types: grey matter (GM), white matter (WM), cerebrospinal fluid (CSF), skull, and scalp. Then, the fixed conductivity values implemented in the SimNIBS were applied for each tissue type. The electric field interpolation was performed using Saturnino et al (95) and computed the electrical field at the centre of GM in the ATL and vertex. Finally, we averaged the individual electrical field for the ATL (Fig. 2C) and vertex (Fig. S6).

Revised version: SimNIBS 3.2 (94) was used to calculate individual electric field of cTBS. The pipeline by Nielsen et al (95) was utilized to generate the individual head model consisting of five tissue types: grey matter (GM), white matter (WM), cerebrospinal fluid (CSF), skull, and scalp. Then, the fixed conductivity values implemented in the SimNIBS were applied for each tissue type. The electric field interpolation was performed using Saturnino et al (96) and computed the electrical field at the centre of GM in the ATL and vertex. Finally, we averaged the individual electrical field for the ATL (Fig. 2C) and vertex (Fig. S6).

Magnetic resonance imaging acquisition

A 3T Philips Achieva MRI scanner was used to acquire data with a 32-channel head coil with a SENSE factor 2.5. Structural images were acquired using a magnetization prepared rapid acquisition gradient echo (MPRAGE) sequence (TR = 8.4 ms, TE = 3.9 ms, slice thickness 0.9 mm, in-planed resolution 0.94 × 0.94 mm).

Earlier version: MRS data were acquired using GABA-edited MEGA-PRESS sequence(96) (TR = 2000ms, TE = 68ms). The voxel of interest (VOI) was manually placed in the left ventrolateral ATL (voxel size = 40 x 20 x 20mm), avoiding hippocampus or vertex (voxel size = 30 x 30 x 30mm), Cz, guided by international 10-20 electrode system (97) (Fig 1.C). Spectra were acquired in interleaved blocks of four scans with application of the MEGA inversion pulses at 1.95 ppm to edit GABA signal (100 repeats at the ATL VOI and 75 repeats at the vertex VOI). Measurements from the ATL VOI with current protocol provided a robust measure of GABA and glx concentrations (48, 49, 98). A total of 1024 sample points were collected at a spectral width of 2 kHz.

Revised version: MRS data were acquired using GABA-edited MEGA-PRESS sequence(97) (TR = 2000ms, TE = 68ms). The voxel of interest (VOI) was manually placed in the left ventrolateral ATL (voxel size = 40 x 20 x 20mm), avoiding hippocampus or vertex (voxel size = 30 x 30 x 30mm), Cz, guided by international 10-20 electrode system (98) (Fig1.C). Spectra were acquired in interleaved blocks of four scans with application of the MEGA inversion pulses at 1.95 ppm to edit GABA signal (100 repeats at the ATL VOI and 75 repeats at the vertex VOI). Measurements from the ATL VOI with current protocol provided a robust measure of GABA and glx concentrations (47, 48, 99). A total of 1024 sample points were collected at a spectral width of 2 kHz.

Earlier version: A dual-echo fMRI protocol developed by Halai et al (99) was employed to maximise signal-to-noise (SNR) in the ATL (TR = 2.8 s, TE = 12 ms and 35 ms, 42 slices, 96 × 96 matrix, 240 × 240 × 126 mm FOV, slice thickness 3 mm, in-plane resolution 2.5 × 2.5).

Revised version: A dual-echo fMRI protocol developed by Halai et al (100) was employed to maximise signal-to-noise (SNR) in the ATL (TR = 2.8 s, TE = 12 ms and 35 ms, 42 slices, 96 × 96 matrix, 240 × 240 × 126 mm FOV, slice thickness 3 mm, in-plane resolution 2.5 × 2.5).

MRS analysis

Earlier version: Java-based magnetic resonance user’s interface (jMRUI5.1, EU project www.jmrui.eu) (100) was used to analyse MRS data. Raw data were corrected using the unsuppressed water signal from the same VOI, eddy current correction, a zero-order phasing of array coil spectra. Residual water was removed using Hankel-Lanczos singular value decomposition (101). Advanced Magnetic Resonance (AMARES) (102) was used to quantify neurochemicals including GABA, glx, and NAA. The exclusion criteria for data were as follows: Cramér-Rao bounds > 50%, water linewidths at full width at half maximum (FWHM) > 20 Hz, and SNR < 40. A subject was discarded from the analysis due to poor quality of MRS. GABA and glx values are reported as a ratio to NAA as we previously reported (48).

Revised version: Java-based magnetic resonance user’s interface (jMRUI5.1, EU project www.jmrui.eu) (101) was used to analyse MRS data. Raw data were corrected using the unsuppressed water signal from the same VOI, eddy current correction, a zero-order phasing of array coil spectra. Residual water was removed using Hankel-Lanczos singular value decomposition (102). Advanced Magnetic Resonance (AMARES) (103) was used to quantify neurochemicals including GABA, glx, and NAA. The exclusion criteria for data were as follows: Cramér-Rao bounds > 50%, water linewidths at full width at half maximum (FWHM) > 20 Hz, and SNR < 40. A subject was discarded from the analysis due to poor quality of MRS. GABA and glx values are reported as a ratio to NAA as we previously reported (47).

Earlier version: Statistical Parametric Map (SPM8, http://www.fil.ion.ucl.ac.uk/spm/) was used to calculate the contributions of GM and WM to the VOI from the structural image. Then voxel registration was performed using custom-made scripts developed in MATLAB by Dr. Nia Goulden, which can be accessed at http://biu.bangor.ac.uk/projects.php.en. The calculation of tissue types within the VOI provided the percentage of each tissue type. As GABA levels are substantially higher (two-fold) in the GM than WM(103), we used GM as a covariate in the analysis. There was no significant difference in GM volume before and after the stimulation (ps > 0.5) and a significant correlation between GM volumes before and after stimulation in both VOIs (ATL stimulation: r = 0.75, p < 0.001 in the ATL, r = 0.67, p = 0.003 in the vertex; Vertex stimulation: r = 0.68, p = 0.008 in the ATL, r = 0.72, p < 0.001). The results of tissue segmentation is summarised in Table S4.

Revised version: Statistical Parametric Map (SPM8, http://www.fil.ion.ucl.ac.uk/spm/) was used to calculate the contributions of GM and WM to the VOI from the structural image. Then voxel registration was performed using custom-made scripts developed in MATLAB by Dr. Nia Goulden, which can be accessed at http://biu.bangor.ac.uk/projects.php.en. The calculation of tissue types within the VOI provided the percentage of each tissue type. As GABA levels are substantially higher (two-fold) in the GM than WM(104), we used GM as a covariate in the analysis. There was no significant difference in GM volume before and after the stimulation (ps > 0.5) and a significant correlation between GM volumes before and after stimulation in both VOIs (ATL stimulation: r = 0.75, p < 0.001 in the ATL, r = 0.67, p = 0.003 in the vertex; Vertex stimulation: r = 0.68, p = 0.008 in the ATL, r = 0.72, p < 0.001). The results of tissue segmentation is summarised in Table S4.

fMRI analysis

Earlier version: fMRI data were processed using SPM8. Dual gradient echo images were realigned, corrected for slice timing, and averaged using in-house MATLAB code developed by Halai et al (99). The EPI volumes were coregistered into the structural image, spatially normalized to the MNI template using DARTEL(diffeomorphic anatomical registration through an exponentiated lie algebra) toolbox (104), and smoothed with an 8 mm full-width half-maximum Gaussian filter.

Revised version: fMRI data were processed using SPM8. Dual gradient echo images were realigned, corrected for slice timing, and averaged using in-house MATLAB code developed by Halai et al (100). The EPI volumes were coregistered into the structural image, spatially normalized to the MNI template using DARTEL(diffeomorphic anatomical registration through an exponentiated lie algebra) toolbox (105), and smoothed with an 8 mm full-width half-maximum Gaussian filter.

A general linear model (GLM) was used to perform statistical analyses. A design matrix was modelled with task conditions, semantic, control and baseline for each individual along with six motion parameters as regressors. A contrast of interest (sematic > control) for each participant were calculated. One-sample t-test was performed to estimate the contrast of interest at the group-level. Clusters were considered significant when passing a threshold of p FWE-corrected < 0.05, with at least 100 contiguous voxels.

Earlier version: Regions of interest (ROI) analysis was conducted using Marsbar (105). The mean signal changes of VOIs were extracted for semantic task condition before and after the stimulation.

Revised version: Regions of interest (ROI) analysis was conducted using Marsbar (106). The mean signal changes of VOIs were extracted for semantic task condition before and after the stimulation.

A voxel-wise simple regression analysis was conducted to identify the local maxima of voxels within the MRS ATL VOI correlating with its BOLD response with GABA levels in the contrast of interest (semantic > control). Local maxima of correlation were estimated on a voxel level, setting the threshold to p < 0.05 FWE after small-volume correction.

Statistical analysis

Earlier version: For behavioural data, accuracy and reaction time (RT) were calculated for each individual. We computed the inverse efficiency (IE) score (RT/1-the proportion of error) to combine the accuracy and RT and calculated the cTBS effect (POST-PRE session). A 2 × 2 repeated measures ANOVA with stimulation (ATL vs. vertex) and task (semantic vs. control) as within-subject factors. Post hoc paired t-tests were conducted.

Revised version: For behavioural data, accuracy and reaction time (RT) were calculated for each individual. We computed the inverse efficiency (IE) score (RT/1-the proportion of error) to combine the accuracy and RT and calculated the cTBS effect (POST-PRE session). A 2 × 2 repeated measures ANOVA was conducted with stimulation (ATL vs. vertex) and task (semantic vs. control) as within-subject factors. Post hoc paired t-tests were conducted.

Earlier version: To investigate the individual-level effects of cTBS, participants were categorized based on changes in their semantic task performance following the ATL stimulation. Individuals exhibiting a decline in task performance post-ATL cTBS in comparison to the pre-stimulation session were classified as responders. Conversely, those showing no changes or an improvement in task performance after ATL cTBS were designated as non-responders. Subsequently, there were 7 responders and 10 non-responders. Planned t-tests were conducted to examine differences between responders and non-responders in semantic task performance and ATL GABA levels in pre-stimulation session. Additionally, planned paired t-tests were performed to investigate the cTBS effects on semantic task performance and ATL GABA levels within each group.

Revised version: To investigate the individual-level effects of cTBS, participants were categorized based on changes in their semantic task performance following the ATL stimulation. Prior to classification, task performance was adjusted to account for practice effects by normalizing ATL stimulation performance relative to vertex stimulation performance. Individuals exhibiting a decline in task performance post-ATL cTBS in comparison to the pre-stimulation session were classified as responders, while those showing no change or improvement in performance were categorized as non-responders. Subsequently, there were 9 responders and 8 non-responders. A 2×2×2 ANOVA was conducted with task (semantic vs. control) and session (PRE vs. POST) as within-subject factors, and group (responders vs. non-responders) as a between-subject factor was performed. Post hoc t-tests were conducted to examine differences in task performance between responders and non-responders. Additionally, a 2 x 2 ANOVA with session (pre vs. post) as a within subject factor and with group (responders vs. non-responders) as a between subject factor was conducted to examine the effects of group in ATL GABA levels. Post hoc t-tests were performed to investigate the cTBS responsiveness on semantic task performance and ATL GABA levels.

Partial correlation analysis was performed to illustrate the relationship between the ATL GABA levels and semantic task performance, accounting for GM volume, age and sex.

Earlier version: Regression analyses (linear and quadratic models) were conducted to explore the relationship between the ATL GABA levels and semantic task performance. The individual ATL GABA levels were adjusted by GM volume, age, and sex, using multiple regression analysis. In order to determine the best-fit model, we calculated the Bayesian Information Criterion (BIC) as a measure of model fitness (51) and performed a partial F-test.

Revised version: Regression analyses (linear and quadratic models) were conducted to explore the relationship between the ATL GABA levels and semantic task performance. The individual ATL GABA levels were adjusted by GM volume, age, and sex, using multiple regression analysis. In order to determine the best-fit model, we calculated the Bayesian Information Criterion (BIC) as a measure of model fitness (50) and performed a partial F-test.

Statistical analyses were undertaken using Statistics Package for the Social Sciences (SPSS, Version 25, IBM Cary, NC, USA) and RStudio (2023).